# Week 1 Study Notebook

# UC San Diego OMDS DSC 207

## Check your Python skills while using Colab notebooks.

### Installing Otter-Grader

Otter-Grader is a powerful tool that helps you check your work before submitting assignments. To get started with Otter-Grader, you need to install it within this Colab notebook.

To install Otter-Grader, simply run the below code cell.

Note that questions 1-3 and 5-9 are autograded, while questions 4 and 10 will be graded manually.

```python
# DO NOT MODIFY
%pip install otter-grader

files = "https://github.com/dsc207rfall2023/fa23course-dsc207r/raw/main/assignments/Week%201/tests.zip"
!wget $files && unzip -o tests.zip
```

Output:
```
Collecting otter-grader
  Downloading otter_grader-6.0.4-py3-none-any.whl.metadata (3.9 kB)
Requirement already satisfied: click<9.0.0,>=8.1.7 in /usr/local/lib/python3.10/dist-packages (from otter-grader) (8.1.8)
Collecting dill>=0.3.0 (from otter-grader)
  Downloading dill-0.3.9-py3-none-any.whl.metadata (10 kB)
Collecting fica>=0.4.1 (from otter-grader)
  Downloading fica-0.4.1-py3-none-any.whl.metadata (2.0 kB)
Collecting ipylab<2.0.0,>=1.0.0 (from otter-grader)
  Downloading ipylab-1.0.0-py3-none-any.whl.metadata (6.7 kB)
Requirement already satisfied: ipython in /usr/local/lib/python3.10/dist-packages (from otter-grader) (7.34.0)
Collecting ipywidgets<9.0.0,>=8.1.5 (from otter-grader)
  Downloading ipywidgets-8.1.5-py3-none-any.whl.metadata (2.3 kB)
Requirement already satisfied: jinja2<4.0,>=3.1 in /usr/local/lib/python3.10/dist-packages (from otter-grader) (3.1.5)
Collecting jupytext<2.0.0,>=1.16.4 (from otter-grader)
  Downloading jupytext-1.16.6-py3-none-any.whl.metadata (13 
```

```python
# Initialize Otter
import otter
grader = otter.Notebook()
```

### Loading Python modules

```python
import pandas as pd
import numpy as np

# ADD ADDITIONAL PYTHON MODULES BELOW
from math import sqrt
```

### 1.  Write a program to swap a string and an int variable and print their data types to verify the swapping operation

* Initialize 2 variables my_int = 5 , and my_str = "Hi"
* Swap these variables such that my_int now holds the value of my_str and vice versa
* Verify this by printing their data types

_Points:_ 5

```python
# Initialize my_int and my_str
my_int = 5
my_str = "Hi"

# swap values for variables my_int and my_str
my_int, my_str = my_str, my_int

# print data type for my_int
print(type(my_int))

# print data type for my_str
print(type(my_str))

```

Output:
```
<class 'str'>
<class 'int'>
```

```python
grader.check("q1")
```

Output:
```
q1 results: All test cases passed!
```

### 2. Write a Python Program to Calculate the length of the diagonal of a rectangle

* Initialize the 2 sides of the rectangle a and b to 5 and 6
* Calculate the length of the diagonal using the formula - diagonal = √(a^2 + b^2)

_Points:_ 5

```python
# Initialize variables a and b
a = 5
b = 6

def calc_diag(a, b):
    diag = sqrt(a**2+b**2)
    return(diag)
```

```python
grader.check("q2")
```

Output:
```
q2 results: All test cases passed!
```

### 3. Write a Python Program to Check if a Number is divisible by 3

* Initialize a variable that takes input from user
* Check and return True or False if the number is divisible by 3
* If the input is not an int, return False

_Points:_ 10

```python
#Initialize a variable that takes input from user
user_input = input('Please enter a number to see if it is divisible by 3: ')

def is_divisible_by_3(num):
    if type(num) == int:
        if num %3 == 0:
            return(True)
        else:
            return(False)
    else:
        return(False)

```

Output:
```
Please enter a number to see if it is divisible by 3: 3
```

```python
grader.check("q3")
```

Output:
```
q3 results: All test cases passed!
```

<!-- BEGIN QUESTION -->

### 4. Write a Python Program to Check if a number has exactly 3 digits (no zeros allowed at the beginning and the number can be assumed to be positive)

* Initialize a variable that takes input from user(input will be given as a string)
* Check and print true if the number has exactly 3 digits(no zeros allowed at the beginning) else print false

_Points:_ 10

```python

'''
Since no zeros are allowed at the beginning I will assume floats are not
included int the test.

Also since there is no test included for this question I will loop through a
list that tests several cases for the num_is_three_digits function.
test_1: positive 3 digit int
test_2: negative 3 digit int
test_3: leading/trailiing whitespace
test_4: number starting with 0
test_5: number with less than 3 digits
test_6: number with greater than 3 digits
test_7: string
'''
num = input('Please enter an integer: ')

tests = ['100','-100',' 100 ','010','1','1000','my_str']

def num_is_three_digits(num):
  # Remove leading or trailing white space
  num = num.strip()
  try:
    # Check if user's number is an int
    if type(int(num)) == int:
      # Remove negative sign
      digits = [digit for digit in list(num) if digit != '-']
      # Check length of user's number
      if len(digits) == 3:
        # Check if number starts with zero
        if digits[0] == '0':
          print('number cannot start with 0')
          return(False)
        else:
          print('number has 3 digits')
          return(True)
      else:
        print(f'number has {len(digits)}')
        return(False)
  except:
    print("not an int please enter int")
    return(False)

result = num_is_three_digits(num)
print(result,'\n')

# Test cases stated above in inital comments
print(80*'=')
print('Test cases below')
print(80*'=')
for test in tests:
  print(f'Testing: {test}')
  result = num_is_three_digits(test)
  print(result)
  print()
```

Output:
```
Please enter an integer: 200
number has 3 digits
True 

================================================================================
Test cases below
================================================================================
Testing: 100
number has 3 digits
True

Testing: -100
number has 3 digits
True

Testing:  100 
number has 3 digits
True

Testing: 010
number cannot start with 0
False

Testing: 1
number has 1
False

Testing: 1000
number has 4
False

Testing: my_str
not an int please enter int
False
```

<!-- END QUESTION -->

### 5. Write a Python Program to Check Whether a String is Palindrome or Not

* Initialize a string
* Lower case it to allow easy comparison
* Check using if-else statements if the string is a palindrome
* A palindrome is a string which is same read forward or backwards.

_Points:_ 10

```python
# Initialize a string
my_str = 'kayak'
my_str = my_str.lower()

def if_palindrome(x):
  if x == x[::-1]:
    return(True)
  else:
    return(False)

```

```python
grader.check("q5")
```

Output:
```
q5 results: All test cases passed!
```

### 6.  **Print numbers divisible by 5 between 4 and 30, inclusive**

* Iterate through the numbers in the for loop and notice the way range has been used to generate numbers in the for loop
* Produce a list and print containing only the numbers divisible by 5 using if statement

_Points:_ 5

```python
start = 4
end = 30

# must add one to the upper bound since the range function is not inclusive
sol = [i for i in range(start,end+1) if i%5 == 0]

```

```python
grader.check("q6")
```

Output:
```
q6 results: All test cases passed!
```

### 7. Write a Python Program to print the first n numbers of the Fibonacci series starting from 0

* Use a loop and return a list of first n numbers of the Fibonacci series: 0,1,1,2,3,....


_Points:_ 15

```python

def fibonacci(n):
    fib_sequence = [0, 1]
    while len(fib_sequence) < n:
      fib_sequence.append(fib_sequence[-1] + fib_sequence[-2])
    return(fib_sequence[0:n])

```

```python
grader.check("q7")
```

Output:
```
q7 results: All test cases passed!
```

### 8. Write a Python Program that defines a function to Remove Vowels From a String

* You can initialize a string of your choice
* Define the function
* Use a for loop to return an updated string with the vowels removed.

_Points:_ 10

```python
def remove_vowels(input_str):
  vowels = list("AEIOUaeiou")
  input_str = [i for i in input_str if i not in vowels]
  input_str = "".join(input_str)
  return(input_str)


```

```python
grader.check("q8")
```

Output:
```
q8 results: All test cases passed!
```

### 9. Write a program to define a function that takes two numbers as arguments and returns the sum of all numbers between the arguments passed (exclusive, i.e. excluding the limits).

* You can define a range of your choice
* Define the function
* Use a for loop to add numbers within the range(exclusive)

_Points:_ 10

```python
def sum(a, b):
    range_of_numbers_exclusive = [i for i in range(a+1,b)]
    sum_of_range = 0
    for i in range_of_numbers_exclusive:
      sum_of_range += i
    return(sum_of_range)

```

```python
grader.check("q9")
```

Output:
```
q9 results: All test cases passed!
```

<!-- BEGIN QUESTION -->

### 10. Define your own Python function!

#### Part A
In Python, functions are blocks of code that only run when they are called. Functions are programmed to complete a specific task. You can input data, called parameters, into functions. Functions can also return data as a result.

Write a function with the following components:

* Pass a list as the function parameter. The list can be composed of any data type (int, string, float, etc.)
* At least one conditional (if/else/elif)
* One loop (ex. for, while, do-while)
* Returns one or more values. This value can be a list or a single element.


Ensure the following:
* Your program contains all the components listed above.
* You test the function by calling it outside of the function.
* Your program runs without any errors.

_Points:_ 10

```python
def remove_strings(my_list):
  new_list=[]
  for item in my_list:
    if type(item) != str:
      new_list.append(item)
    else:
      print(f"do not append this item: {item}")
  return(new_list)

test_list = ['a',True,1,2.34,'my_str']
removed_strings_list = remove_strings(test_list)
print(removed_strings_list)

```

Output:
```
do not append this item: a
do not append this item: my_str
[True, 1, 2.34]
```

#### Part B
Record a video, no more than 3 minutes long.

* Explain the function at a high level. We want you to explain the flow of code, including the order in which the lines of code are executed.
* Explain at least two of the required components. For example, you could describe your conditional statement and what your loop does.

You can reference the example video here: https://drive.google.com/file/d/1LragCuaHO_eVMm1EDfFv7wqQcrC86NuW/view?usp=sharing

Upload the recorded video to your drive and paste the shareable link in the cell below.

Note: Ensure that we are able to view the video.

_Points:_ 10

```python
#https://drive.google.com/file/d/1NoM7Vm1Fv7DvJ65aMJp6WI-5aeHTIxxq/view?usp=drive_link
```

<!-- END QUESTION -->

### Submission

Before you submit your notebook, ensure that you've run all the cells sequentially to display images, graphs, and outputs correctly (if the outputs are missing we might not be able to award points for those parts). Take the time to review your solutions and evaluate using the public test cases. Once you're satisfied, save your notebook by navigating to "File" in the Colab Notebook menu and selecting "Save" or using the keyboard shortcut (usually Ctrl + S or Cmd + S).

To submit, export your notebook as an .ipynb file. Ensure it's saved with the same name as the assignment. Then, upload this .ipynb file to Gradescope. Remember to submit the correct version.

Please save your work before exporting it, and if you encounter any technical issues or have questions about the submission process, reach out to the course staff for assistance.

---

To double-check your work for questions 1-3 and 5-9, the cell below will rerun all of the autograder tests.

```python
grader.check_all()
```

Output:
```
q1 results:
    q1 - 1 result:
        ✅ Test case passed

    q1 - 2 result:
        ❌ Test case failed
        Trying:
            assert type(my_str) == int
        Expecting nothing
        **********************************************************************
        Line 1, in q1 1
        Failed example:
            assert type(my_str) == int
        Exception raised:
            Traceback (most recent call last):
              File "/usr/lib/python3.10/doctest.py", line 1350, in __run
                exec(compile(example.source, filename, "single",
              File "<doctest q1 1[0]>", line 1, in <module>
                assert type(my_str) == int
            AssertionError

    q1 - 3 result:
        ✅ Test case passed

    q1 - 4 result:
        ❌ Test case failed
        Trying:
            assert my_str == 5
        Expecting nothing
        **********************************************************************
        Line 1, in q1 3
        Failed example:
            a
```

Tests for q1 (points: 5)
```
>>> assert type(my_int) == str
```
```
>>> assert type(my_str) == int
```
```
>>> assert my_int == "Hi"
```
```
>>> assert my_str == 5
```


Tests for q2 (points: 5)
```
>>> assert calc_diag(3, 4) == 5
```


Tests for q3 (points: 10)
```
>>> assert is_divisible_by_3(9) == True
```
```
>>> assert is_divisible_by_3(-12) == True
```


Tests for q5 (points: 10)
```
>>> assert if_palindrome("racecar") == True
```
```
>>> assert if_palindrome("hello") == False
```


Tests for q6 (points: 5)
```
>>> assert len(sol) == 6
```
```
>>> assert sol == [5, 10, 15, 20, 25, 30]
>>>
```


Tests for q7 (points: 15)
```
>>> assert fibonacci(1) == [0]
```
```
>>> assert fibonacci(5) == [0, 1, 1, 2, 3]
```


Tests for q8 (points: 10)
```
>>> assert remove_vowels("Hello, World!") == "Hll, Wrld!"
```
```
>>> assert remove_vowels("Python Programming") == "Pythn Prgrmmng"
```


Tests for q9 (points: 10)
```
>>> assert sum(1, 5) == 9
>>>
```
```
>>> assert sum(-3, 3) == 0
```
